<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

# Aufgabenblatt 1

## Aufgabe 1: Formatierung

Warum benötigt z. B. eine for-Schleife der Form

```python
for i in range (1, 10):
    print (i)
    print (i, "mal 2 ist", i*2)
```
keine geschweiften Klammern, Schlüsselwörter begin und end oder andere syntaktische Hilfsmittel, um zu erkennen, welche Befehle innerhalb der Schleife ausgeführt werden sollen?

*Tragen Sie hier Ihre Lösung ein*

## Aufgabe 2: Bash-Skript in Python Skript umwandeln

Wandeln Sie das folgende Shell-Skript, das die Fakultät $i!$ für die Werte $𝑖 \in \{1, 2, \dots, 20\}$ berechnet, in ein Python-Skript um:

```bash
for i in $( seq 1 20 ); do
  count=$i; ergebnis=1
  while [ $count -gt 1 ]; do
    ergebnis=$((ergebnis * count))
    count=$((count-1))
  done
  echo "fakultaet" $i "=" $ergebnis
done
```

<!-- grade_id: cell-c3ed84867410823a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-026b282fd4ee3bd5, points: 1, grade: true, solution: false -->
```python
assert i == 20, 'range function does not include the upper limit'
assert ergebnis == 243290200817664e4, 'factorial of 20 is incorrect'
assert count == 1, 'count should reduce to 1'
if '%(' not in _i.replace(' ',''):
    print('Pass your variables in a tuple to print statement')
```

## Aufgabe 3: Gleichheit und Identität

1. Testen Sie die beiden Vergleichsoperatoren `==` (Gleichheit) und `is` (Identität) mit folgenden Befehlen:

```python
a = 3
b = 3
a == b # Test auf Gleichheit, True oder False?
a is b # Test auf Identität, True oder False?

c = [3,4]
d = [3,4]
e = d
c == d
d == e
c is d
d is e
```

<!-- grade_id: cell-a2a84a461a150adc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-af246e1bcd365cd4, points: 0, grade: true, solution: false -->
```python
assert type(a) == int
assert type(a) == int
assert a is b, 'Try smaller integers'
```

2. Sie haben das Ergebnis von `a is b` für den Fall `a=3` und `b=3` gesehen. Wiederholen Sie den Test für größere Werte von `a` und `b`. Fällt das Ergbenist von `a is b` immer gleich aus?

<!-- grade_id: cell-516ae13f0714dcc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e2d206b5e6776880, points: 1, grade: true, solution: false -->
```python
assert a == b, 'assign the same integer'
assert a > 256, 'choose larger number'
```

## Aufgabe 4: Divisionsoperator

Wie unterscheiden sich die beiden Operatoren `/` und `//`?

Rechnen Sie dazu mit Integern `5/3` und `5//3` sowie mit Float-Zahlen `5.0/3.0` und `5.0//3.0` aus.

```python

```

**Achtung:** Der Divisionsoperator `/` hat eine unterschiedliche Bedeutung in Python2.
Dort ergibt `5/3 = 1`.

Manchmal sehen Sie zu Begin eines Python Codes folgendes:
```python
from __future__ import division
```
Damit kann man für einzelne Aspekte (z.B. die Division) in Python2 das gleiche Verhalten erreichen wie in Python3.

```python

```

## Aufgabe 5: Zahlensysteme

Für die häufig benötigten Zahlensysteme mit Basis 2, 8 und 16 können Sie in Python Zahlen mit den Präfixen `0b` (binär, Basis 2), `0o` (oktal, Basis 8) und `0x` (hexadezimal, Basis 16) darstellen (z. B. `0xa = 0o12 = 0b1010 = 10`). Alternativ verwenden Sie die int-Funktion, die einen String und die gewünschte Zahlenbasis als Parameter akzeptiert (z. B. `int("a",16)` oder `int("1010",2)` für hexadezimale und binäre Zahlen).

Lassen Sie sich mit beiden Methoden (Präfix bzw. int-Funktion) die Dezimaldarstellungen der folgenden Zahlen ausgeben:
- `10010101` (binär)
- `7777` (oktal)
- `c000` (hexadezimal)

<!-- grade_id: cell-4e845cd6f5cc5e66, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd3d248d37627184, points: 1, grade: true, solution: false -->
```python
assert ('0b' or '0o' or '0x') in _i, 'Use one of three prefix'
assert type(_) is int, 'cell output should be integer'
```

```python
int("10010101",2)
```

<!-- grade_id: cell-4bd6c1a7445e8e29, points: 1, grade: true, solution: false -->
```python
assert 'int(' in _i, 'use int function'
assert type(_) is int, 'cell output should be integer'
```

Umgekehrt können Sie über Format-Strings mit `%o` und `%x` Zahlen im Oktal- und Hexadezimalformat ausgeben. Testen Sie z. B.
die folgenden Befehle:
- `"%x" % 49152`
- `"%o" % 18056`

```python

```

Der Einsatz des Operators `%` gilt in Python 3 allerdings als veraltet – stattdessen wird hier die auf Strings anwendbare Methode `format()` verwendet, z. B. wie folgt:
```
formatstring = "dezimal: {}, hex.: {:x}, oktal: {:o}, binär: {:b}"
w = 99
formatstring.format (w,w,w,w)
```

Rufen Sie mit `help('FORMATTING')` die eingebaute Hilfe zur Formatierung auf und schauen Sie sich die Syntax und die am Ende
gezeigten Beispiele an.

```python

```

# Aufgabe 6: Schleifen

Programmieren Sie eine geschachtelte Schleife, die unten stehende Tabelle erzeugt.

|   x   |   y   |   z   | x and y or ((y and z) and not x) |
| ----- | ----- | ----- | -------------------- |
| True  | True  | True  | True| 
| False  | True  | True  | True| 
| True  | False  | True  | False| 
| False  | False  | True  | False| 
| True  | True  | False  | True| 
| False  | True  | False  | False| 
| True  | False  | False  | False| 
| False  | False  | False  | False | 


<!-- grade_id: cell-9de9edb007024bca, points: 0, grade: false, solution: true -->
```python
print("x", "\t", "y", "\t", "z", "\t", "x and y or ((y and z) and not x)")
print('---------------------------------------------------------')
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture screen --no-stderr --no-display
# rerunning the code in the last cell to capture the output
c = _i
exec(c)
```

<!-- grade_id: cell-2e1e33745a47e24e, points: 1, grade: true, solution: false -->
```python
assert c.count('for')>2, 'Use nested loops'

# str is converted to bool
s = screen.stdout.split('\n')
s = [[i == 'True' for i in line.split()] for line in s[2:-1]]
for i in s:
    if ((i[0] and i[1]) or ((i[1] and i[2]) and not i[0])) != i[3]:
        print('Results are not consistent!')
        raise

assert screen.stdout.count('\t') > 24, 'Use \\t to seperate values'
```

## Aufgabe 7: Dateien

Schreiben Sie mehrere Strings als Zeilen in eine Datei und schließen Sie die Datei.

<!-- grade_id: cell-98b3d0ad08dea97e, points: 0, grade: false, solution: true -->
```python
file_name = "ausgabe.txt"

fobj = None #Use this variable for your file
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c6e009aa21f30a9b, points: 1, grade: true, solution: false -->
```python
assert fobj.closed, 'File is not closed'
try:
    file = open(file_name, "r")
    assert file.read().count('\n')>2, '"ausgabe.txt" does not contain at least three lines'
    file.close()
except FileNotFoundError:
    print('"ausgabe.txt" has not been created')
    raise
```

Öffnen Sie die Datei zum Lesen und geben Sie den Inhalt aus.

<!-- grade_id: cell-07b94271a607a39c, points: 0, grade: false, solution: true -->
```python
fobj = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture screen --no-stderr --no-display
# rerunning the code in the last cell to capture the output
c = _i
exec(c)
```

<!-- grade_id: cell-ec00e685159461d9, points: 1, grade: true, solution: false -->
```python
assert fobj.closed, 'File is not closed'
assert screen.stdout.count('\n')>2, 'Not all lines are printed'
```

## Aufgabe 8: Funktionen

Geben Sie die folgende Funktion mit unendlicher Rekursion in ipython ein und testen Sie sie durch Aufruf von `unendlich(0)`.
```python
def unendlich(n):
    return unendlich(n+1)
```

<!-- grade_id: cell-63f20cc0d8f4b0b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Die Fibonacci-Funktion ist durch

$\large
fib(n) = 
\begin{cases}
    0,& \text{if } n = 0\\
    1,& \text{if } n = 1\\
    fib (n–1) + fib (n–2), & \text{if } n \geq 2
\end{cases}$

definiert. Betrachten Sie die folgenden zwei Implementierungen von `fib`:
```python
def fib_a(n):
    if n == 0:
        return 0
    elif n == 1:
        return 1
    else:
        return fib_a(n-1) + fib_a(n-2)
```

```python
def fib_b(n):
    (a,b) = (0,1)
    for i in range(n):
        (a,b) = (b,a+b)
    return a
```

1. Welche der beiden Funktionen ist rekursiv?
2. Geben Sie beide Funktionen in eine Code-Zelle ein und rufen Sie sie mit den Argumenten 20, 30 und 35 auf. Sie werden sehr unterschiedliche Laufzeiten beobachten. Haben Sie dafür eine Erklärung?

<!-- grade_id: cell-0455db3dc971ae3a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d667e3a0b5606420, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

3. Schreiben Sie eine bessere rekursive Implementierung von fib, die ein Dictionary verwendet,
welches als
```python
fib_dict = {}
fib_dict[0] = 0
fib_dict[1] = 1
```
initialisiert wird. Die neue Funktion fib_c() soll prüfen, ob der gesuchte Wert bereits im Dictionary enthalten ist – wenn ja, kann sie ihn einfach zurückgeben. (Zur Erinnerung: Ob ein Schlüssel / Key im Dictionary enthalten ist, testen Sie mit `key in dict`.) Wenn nicht, soll sie ihn rekursiv, also durch Aufruf von fib_c(), berechnen.
4. Das Dictionary fib_dict deklarieren Sie außerhalb der Funktion. Warum ist das
vorteilhaft?

<!-- grade_id: cell-9a396222d44ded2f, points: 0, grade: false, solution: true -->
```python
fib_dict = {}
fib_dict[0] = 0
fib_dict[1] = 1

def fib_c(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2dcea799bd9d4a77, points: 1, grade: true, solution: false -->
```python
try:
    fib_c(0)
    fib_c(1)
except:
    print('add cases of 0 and 1')
    raise

test = [(0,0),(1,1),(10, 55),(20, 6765),(30, 832040)]
assert not [a for a, b in test if fib_c(a) != b], 'Results are incorrect!'
fib_c(30)
```

## Aufgabe 9: Strings und Konvertierung

Implementieren Sie eine Funktion `sternchen()`, die als Argument einen String oder eine Zahl akzeptiert. Sie soll als Rückgabewert einen String liefern, der vorne und hinten je drei Sternchen und dazwischen, durch Leerzeichen abgetrennt, den Parameter enthält. Also z. B.:
```python
sternchen("Hallo") = "*** Hallo ***"
sternchen(42) = "*** 42 ***"
sternchen(4.123) = "*** 4.123 ***"
```
Sie können für die Umwandlung von Zahlen in Strings die Funktion `str()` verwenden, sollen diese aber nur aufrufen, wenn der Parameter kein String ist. Bei der Implementierung wird Ihnen die Funktion `type()` helfen.

<!-- grade_id: cell-eea9ddfde999dfcd, points: 0, grade: false, solution: true -->
```python
def sternchen(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a3fff1ffc95e35e0, points: 1, grade: true, solution: false -->
```python
assert not [a for a in [15, 'Yes'] if sternchen(a) != f'*** {a} ***'], 'consider converting integers'
sternchen("HAllo")
```

```python

```