# Aufgabenblatt 1

```python
!pip install --user nose
```

Output:
```
Collecting nose
  Using cached https://files.pythonhosted.org/packages/15/d8/dd071918c040f50fa1cf80da16423af51ff8ce4a0f2399b7bf8de45ac3d9/nose-1.3.7-py3-none-any.whl
Installing collected packages: nose
[33m  WARNING: The scripts nosetests and nosetests-3.4 are installed in '/home/jupyter-giefers.heiner/.local/bin' which is not on PATH.
  Consider adding this directory to PATH or, if you prefer to suppress this warning, use --no-warn-script-location.[0m
Successfully installed nose-1.3.7
```

```python
!ls
```

Output:
```
ex00.ipynb  ex01.ipynb	ipython.png  notebooks.png
```

## Aufgabe 1: Formatierung

Warum benötigt z. B. eine for-Schleife der Form

```python
for i in range (1, 10):
    print (i)
    print (i, "mal 2 ist", i*2)
```
keine geschweiften Klammern, Schlüsselwörter `begin` und `end` oder andere syntaktische Hilfsmittel, um zu erkennen, welche Befehle innerhalb der Schleife ausgeführt werden sollen?

## Aufgabe 2: Bash-Skript in Python Skript umwandeln

Wandeln Sie das folgende Shell-Skript, das die Fakultät $i!$ für die Werte $𝑖 \in \{1, 2, \dots, 20\}$ berechnet, in ein Python-Skript um:

```bash
for i in $( seq 1 20 ); do
  count=$i; ergebnis=1
  while [ $count -gt 1 ]; do
    ergebnis=$((ergebnis * count))
    count=$((count-1))
  done
  echo "fakultaet" $i "=" $ergebnis
done
```

<!-- grade_id: cell-bcde31976cd6b7a0, points: 0, grade: false, solution: true -->
```python
for i in range(1, 21):
    count=i
    ergebnis=1
    while count > 1:
        ergebnis=ergebnis * count
        count=count-1
    print("fakultaet %2d = %20d" % (i, ergebnis))
```

Output:
```
fakultaet  1 =                    1
fakultaet  2 =                    2
fakultaet  3 =                    6
fakultaet  4 =                   24
fakultaet  5 =                  120
fakultaet  6 =                  720
fakultaet  7 =                 5040
fakultaet  8 =                40320
fakultaet  9 =               362880
fakultaet 10 =              3628800
fakultaet 11 =             39916800
fakultaet 12 =            479001600
fakultaet 13 =           6227020800
fakultaet 14 =          87178291200
fakultaet 15 =        1307674368000
fakultaet 16 =       20922789888000
fakultaet 17 =      355687428096000
fakultaet 18 =     6402373705728000
fakultaet 19 =   121645100408832000
fakultaet 20 =  2432902008176640000
```

<!-- grade_id: cell-181c1e65c707c55d, points: 0, grade: true, solution: false -->
```python
if '%(' not in _i.replace(' ',''):
    print('Pass your variables in a tuple to print statement')
```

<!-- grade_id: cell-afff1b8ee835597a, points: 1, grade: true, solution: false -->
```python
assert i == 20, 'range function does not include the upper limit'
assert ergebnis == 243290200817664e4, 'factorial of 20 is incorrect'
assert count == 1, 'count should reduce to 1'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-8-87c0a3b7779b> in <module>
      1 assert i == 20, 'range function does not include the upper limit'
----> 2 assert ergebnis == 243290200817664e4, 'factorial of 20 is incorrect'
      3 assert count == 1, 'count should reduce to 1'

AssertionError: factorial of 20 is incorrect
```

## Aufgabe 3: Gleichheit und Identität

1. Testen Sie die beiden Vergleichsoperatoren `==` (Gleichheit) und `is` (Identität) mit folgenden Befehlen:

```python
a = 3
b = 3
a == b # Test auf Gleichheit, True oder False?
a is b # Test auf Identität, True oder False?

c = [3,4]
d = [3,4]
e = d
c == d
d == e
c is d
d is e
```

<!-- grade_id: cell-67520b0215c0edde, points: 0, grade: false, solution: true -->
```python
a = 3
b = 3
```

<!-- grade_id: cell-33b8663631304fb6, points: 0, grade: true, solution: false -->
```python
assert type(a) == int
assert type(a) == int
assert a is b, 'Try smaller integers'
```

2. Sie haben das Ergebnis von `a is b` für den Fall `a=3` und `b=3` gesehen. Wiederholen Sie den Test für größere Werte von `a` und `b`. Fällt das Ergebnis von `a is b` immer gleich aus?

<!-- grade_id: cell-05aa28e9ca5d755c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8bcb82eba24b8c95, points: 1, grade: true, solution: false -->
```python
assert a == b, 'assign the same integer'
assert a > 256, 'choose larger number'
```

## Aufgabe 4: Divisionsoperator

Wie unterscheiden sich die beiden Operatoren `/` und `//`?

Rechnen Sie dazu mit Integern `5/3` und `5//3` sowie mit Float-Zahlen `5.0/3.0` und `5.0//3.0` aus.

```python

```

**Achtung:** Der Divisionsoperator `/` hat eine unterschiedliche Bedeutung in Python2.
Dort ergibt `5/3 = 1`.

Manchmal sehen Sie zu Begin eines Python Codes folgendes:
```python
from __future__ import division
```
Damit kann man für einzelne Aspekte (z.B. die Division) in Python2 das gleiche Verhalten erreichen wie in Python3.

## Aufgabe 5: Zahlensysteme

Für die häufig benötigten Zahlensysteme mit Basis 2, 8 und 16 können Sie in Python Zahlen mit den Präfixen `0b` (binär, Basis 2), `0o` (oktal, Basis 8) und `0x` (hexadezimal, Basis 16) darstellen (z. B. `0xa = 0o12 = 0b1010 = 10`). Alternativ verwenden Sie die int-Funktion, die einen String und die gewünschte Zahlenbasis als Parameter akzeptiert (z. B. `int("a",16)` oder `int("1010",2)` für hexadezimale und binäre Zahlen).

Lassen Sie sich mit beiden Methoden (Präfix bzw. int-Funktion) die Dezimaldarstellungen der folgenden Zahlen ausgeben:
- `10010101` (binär)
- `7777` (oktal)
- `c000` (hexadezimal)

<!-- grade_id: cell-959f50da2c18a4a6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7e8dc178559ed38b, points: 0, grade: true, solution: false -->
```python
assert ('0b' or '0o' or '0x') in _i, 'Use one of three prefix'
assert type(_) is int, 'cell output should be integer'
```

<!-- grade_id: cell-dc7cd373b78cd8cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-66a2150df1a514ac, points: 0, grade: true, solution: false -->
```python
assert 'int(' in _i, 'use int function'
assert type(_) is int, 'cell output should be integer'
```

Umgekehrt können Sie über Format-Strings mit `%o` und `%x` Zahlen im Oktal- und Hexadezimalformat ausgeben. Testen Sie z. B.
die folgenden Befehle:
- `"%x" % 49152`
- `"%o" % 18056`

```python

```

Der Einsatz des Operators `%` gilt in Python 3 allerdings als veraltet – stattdessen wird hier die auf Strings anwendbare Methode `format()` verwendet, z. B. wie folgt:
```
formatstring = "dezimal: {}, hex.: {:x}, oktal: {:o}, binär: {:b}"
w = 99
formatstring.format (w,w,w,w)
```

Rufen Sie mit `help('FORMATTING')` die eingebaute Hilfe zur Formatierung auf und schauen Sie sich die Syntax und die am Ende
gezeigten Beispiele an.

```python

```

# Aufgabe 6: Schleifen

Programmieren Sie eine geschachtelte Schleife, die unten stehende Tabelle erzeugt.

|   x   |   y   |   z   | x and y or ((y and z) and not x) |
| ----- | ----- | ----- | -------------------- |
| True  | True  | True  | True| 
| False  | True  | True  | True| 
| True  | False  | True  | False| 
| False  | False  | True  | False| 
| True  | True  | False  | True| 
| False  | True  | False  | False| 
| True  | False  | False  | False| 
| False  | False  | False  | False | 


<!-- grade_id: cell-0e62f68a02ffb2fc, points: 0, grade: false, solution: true -->
```python
print("x", "\t", "y", "\t", "z", "\t", "x and y or ((y and z) and not x)")
print('---------------------------------------------------------')
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-47afdd9de01e6374, points: 0, grade: true, solution: false -->
```python
assert _i.count('for')>2, 'Use nested loops'
```

## Aufgabe 7: Dateien

Schreiben Sie mehrere Strings als Zeilen in eine Datei und schließen Sie die Datei.

<!-- grade_id: cell-cf41fcf00f6e7896, points: 0, grade: false, solution: true -->
```python
file_name = "ausgabe.txt"

fobj = None #Use this variable for your file
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-468281cf0db2cad3, points: 1, grade: true, solution: false -->
```python
assert fobj.closed, 'File is not closed'
try:
    file = open(file_name, "r")
    assert file.read().count('\n')>2, '"ausgabe.txt" does not contain at least three lines'
    file.close()
except FileNotFoundError:
    print('"ausgabe.txt" has not been created')
    raise
```

Öffnen Sie die Datei zum Lesen und geben Sie den Inhalt aus.

<!-- grade_id: cell-9d3099a6ab399185, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7b82d8c7638ff041, points: 1, grade: true, solution: false -->
```python
assert fobj.closed, 'File is not closed'

```

## Aufgabe 8: Funktionen

Geben Sie die folgende Funktion mit unendlicher Rekursion in ipython ein und testen Sie sie durch Aufruf von `unendlich(0)`.
```python
def unendlich(n):
    return unendlich(n+1)
```

```python
count = 0;
def unendlich(n):
    global count
    count += 1
    return unendlich(n+1)

try:
    unendlich(0)
except RecursionError:
    print(count)
```

Wie Sie sehen, ist die maximale anzahl der Aufrufe begrenzt; der Aufruf `unendelich(0)` bricht nach einer Reihe von rekursiven Aufrufen mit einem Fehler ab.
Damit an dieser Stelle der auftretende Fehler nicht zum Abbruch der Programmzelle führt, haben wir den Aufruf in einen `try`-`except` Block egstellt.
Um mitzuzählen, wie viele Aufrufe erfolgen, verwenden wir eine globale Variable.

Widmen wir uns nun *dem* klassischen Beispiel, wenn es um Rekursion geht:
Die **Fibonacci-Funktion** ist durch

$\large
fib(n) = 
\begin{cases}
    0,& \text{if } n = 0\\
    1,& \text{if } n = 1\\
    fib (n–1) + fib (n–2), & \text{if } n \geq 2
\end{cases}$

definiert. Betrachten Sie die folgenden zwei Implementierungen von `fib`:
```python
def fib_a(n):
    if n == 0:
        return 0
    elif n == 1:
        return 1
    else:
        return fib_a(n-1) + fib_a(n-2)
```

```python
def fib_b(n):
    (a,b) = (0,1)
    for i in range(n):
        (a,b) = (b,a+b)
    return a
```

1. Welche der beiden Funktionen ist rekursiv?
2. Geben Sie beide Funktionen in eine Code-Zelle ein und rufen Sie sie mit den Argumenten 20, 30 und 35 auf. Sie werden sehr unterschiedliche Laufzeiten beobachten. Haben Sie dafür eine Erklärung?

<!-- grade_id: cell-b461b2a96c2c33cd, points: 0, grade: false, solution: true -->
```python
def fib_a(n):
    if n == 0:
        return 0
    elif n == 1:
        return 1
    else:
        return fib_a(n-1) + fib_a(n-2)
    
def test_fib_a():
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-cd834efb521a5f20, points: 1, grade: true, solution: false -->
```python
from unittest.mock import patch, call
with patch('__main__.fib_a') as mock_f:
     test_fib_a()
calls = [call(20),call(30),call(35)]
mock_f.assert_has_calls(calls, any_order=True)
```

<!-- grade_id: cell-1b47383e77a06e3c, points: 0, grade: false, solution: true -->
```python
def fib_b(n):
    (a,b) = (0,1)
    for i in range(n):
        (a,b) = (b,a+b)
    return a

def test_fib_b():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-db06a701be5a4838, points: 1, grade: true, solution: false -->
```python
from unittest.mock import patch, call
with patch('__main__.fib_b') as mock_f:
     test_fib_b()
calls = [call(20),call(30),call(35)]
mock_f.assert_has_calls(calls, any_order=True)
```

3. Schreiben Sie eine bessere rekursive Implementierung von fib, die ein Dictionary verwendet,
welches als
```python
fib_dict = {}
fib_dict[0] = 0
fib_dict[1] = 1
```
initialisiert wird. Die neue Funktion fib_c() soll prüfen, ob der gesuchte Wert bereits im Dictionary enthalten ist – wenn ja, kann sie ihn einfach zurückgeben. (Zur Erinnerung: Ob ein Schlüssel / Key im Dictionary enthalten ist, testen Sie mit `key in dict`.) Wenn nicht, soll sie ihn rekursiv, also durch Aufruf von fib_c(), berechnen.
4. Das Dictionary fib_dict deklarieren Sie außerhalb der Funktion. Warum ist das
vorteilhaft?

<!-- grade_id: cell-0fd97e428251a494, points: 0, grade: false, solution: true -->
```python
fib_dict = {}
fib_dict[0] = 0
fib_dict[1] = 1

def fib_c(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a8684bbfcb6af051, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

try:
    fib_c(0)
    fib_c(1)
except:
    print('add cases of 0 and 1')
    raise


assert_equal(fib_c(0), 0)
assert_equal(fib_c(1), 1)
assert_equal(fib_c(10), 55)
assert_equal(fib_c(20), 6765)
assert_equal(fib_c(30), 832040)

```

## Aufgabe 9: Strings und Konvertierung

Implementieren Sie eine Funktion `sternchen()`, die als Argument einen String oder eine Zahl akzeptiert. Sie soll als Rückgabewert einen String liefern, der vorne und hinten je drei Sternchen und dazwischen, durch Leerzeichen abgetrennt, den Parameter enthält. Also z. B.:
```python
sternchen("Hallo") # soll "*** Hallo ***" zurückgeben
sternchen(42)      # soll "*** 42 ***"    zurückgeben
sternchen(4.123)   # soll "*** 4.123 ***" zurückgeben
```
Sie können für die Umwandlung von Zahlen in Strings die Funktion `str()` verwenden, sollen diese aber nur aufrufen, wenn der Parameter kein String ist. Bei der Implementierung wird Ihnen die Funktion `type()` helfen.

<!-- grade_id: cell-531e03f36151f79d, points: 0, grade: false, solution: true -->
```python
def sternchen(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-caadc71deef8af87, points: 1, grade: true, solution: false -->
```python
assert not [a for a in [15, 'Yes'] if sternchen(a) != f'*** {a} ***'], 'consider converting integers'
sternchen("HAllo")
```

```python

```