<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

# Aufgabenblatt 2: Galgenmännchen

## Spielidee

Galgenmännchen (im Englischen *hangman game*) spielt man normalerweise mit Papier und Stift.
Normalerweise spielt man Galgenmännchen zu zweit: Der erste Spieler überlegt sich ein Wort und gibt die Anzahl der Buchstaben des Wortes bekannt.
Meist passiert das, indem man entsprechend viele (Unter-) Striche auf das Papier Schreibt.
Der zweite Spieler beginnt nun, Buchstaben zu raten.
Ist der Buchstabe im Wort enthalten, so schreibt Spieler Eins ihn an die entsprechende Stelle auf dem Papier.
Ist der Buchstabe nicht im Wort enthalten, so zeichnet Spieler eins ein weiteres Element des *Galgenmännchens*.
Errät Spieler zwei das Wort, bevor das Galgenmännchen fertig gezeichnet ist, gewinnt er, ansonsten gewinnt Spieler eins.

Ein Programm für das Galgenmännchen Spiel sollte aus folgenden Schritten bestehen:

1. Wähle ein zufälliges Wort (aus einer Liste von Wörtern) aus
2. Begrüße die Nutzerin bzw. der Nutzer
3. Stelle das Spielfeld dar
4. Nimm einen geratenen Buchstaben vom User entgegen
5. Überprüfe, ob der Buchstabe im Wort enthalten ist. Wenn ja, zeige das Wort entsprechend an. Wenn nicht, stelle die nächste Stufe des Galgenmännchens dar
6. Gehe zu Schritt 4, solange das Galgenmännchen noch nicht fertig ist
7. Frage die Nutzerin bzw. der Nutzer, ob sie oder er erneut spielen möchte

## Helferfunktionen

1. Verwende `take_guess` um einen Buchstaben von der Nutzerin bzw. dem Nutzer entgegenzunehmen. Das Spiel wird ausschließlich Großbuchstaben verwenden.
2. Verwende `random_word` um ein zufälliges Wort aus euner Liste auszwählen
3. In der Liste `stages` sind die verschiedenen Phasen des Galgenmännchens abgespeichert. An dieser Stelle helfen uns die *multi-line* Strings mit den drei aufeinanderfolgenden Anführungszeichen.

```python
import random
from io import StringIO
from unittest.mock import patch

def take_guess():
    """catches a single character from the player."""
    
    c = input('Please, guess a letter: ')
    return c.upper() if len(c) == 1 else print('Invalid guess. Try again!')

def random_word():
    """returns a random word."""
    # feel free to change the words in the list
    word_list = ['earthquake', 'quarter', 'nation', 'coast', 'suggestion', 'tub',
                 'smell', 'pigs', 'wilderness', 'part', 'snails', 'twist']
    return random.choice(word_list).upper()

# stages of the hangman
stages = [
        '''
            --------
            |      |
            |      
            |    
            |      
            |     
            -
        ''',
        '''
            --------
            |      |
            |      O
            |    
            |      
            |     
            -
        ''',
        '''
            --------
            |      |
            |      O
            |      |
            |      |
            |     
            -
        ''',
        '''
            --------
            |      |
            |      O
            |     \|
            |      |
            |     
            -
        ''',
        '''
            --------
            |      |
            |      O
            |     \|/
            |      |
            |     
            -
        ''',
        '''
            --------
            |      |
            |      O
            |     \|/
            |      |
            |     /
            -
        ''',
        '''
            --------
            |      |
            |      O
            |     \|/
            |      |
            |     / \\
            -
        '''
    ]
```

## Das Spiel

### Darstellung des Ratewortes

Schreibe eine Funktion die ein `word` als String und eine Liste von geratenen Buchstaben `tru_guess` als Argumente entgegennimmt und als Resultat einen String zurückliefert, bei dem alle nicht-geratenen Buchstaben, als Unterstriche dargestellt sind.

```
'W _ _ D'
```

<!-- grade_id: cell-e0695da39c5d7185, points: 0, grade: false, solution: true -->
```python
def dis_word(word, tru_guess = []):
    """returns a string with only the true guesses visible"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4cc3e96a778a43d6, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------
assert dis_word('Book') == '_ _ _ _', 'For empty true guess list, 4 spaced underscores should be returned.'
assert dis_word('BOOK', 'O') == '_ O O _', 'Elements of true guess list only should appear in the resultant word.'
```

### Galgenmännchen anzeigen

Schreibe eine Funktion, die das Galgenmännchen in ser aktuellen Phase anzeigt. Der Parameter `attempt` gibt dabei an, wie viele falsche Buchstaben bereits geraten wurden. Als Rückgabe soll die Funktion ein String Objekt mit dem aktuellen Hangman zurückliefern.

<!-- grade_id: cell-70ed6d4e27744733, points: 0, grade: false, solution: true -->
```python
def dis_hangman(attempt):
    """returns the correct hangman proper stage for the number of attempts."""
    
    # Return the appropriate item from `stages`
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e614239d362eb0a8, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------
assert dis_hangman(2).replace(' ','') == '\n--------\n||\n|O\n||\n||\n|\n-\n'
```

### Spielstand ausgeben

Schreibe eine Funktion, die auf Grundlage der Parameter `word`, `attempt` und `tru_guess` die Darstellung des Ratewortes und des Galgenmännchens ausgibt.

<!-- grade_id: cell-ad70a29a5fb918f3, points: 0, grade: false, solution: true -->
```python
def dis_state(word, attempt, tru_guess):
    """prints the game state."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-72e893da14d31df4, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------
with patch('sys.stdout', new_callable=StringIO) as screen:
    dis_state('WORT', 2, ['O'])
    for n in [dis_hangman(2), dis_word('WORT', 'O')]:
        assert n in screen.getvalue(), 'dis_hangman and dis_word are not called!'
```

### Spielablauf

Schreiben Sie eine Funktion für den Spielablauf:

1. Verwende die Funktion `random_word` um das Zufallswort für die neue Runde zu erhalten. Inittialisiere die Attribute `attempt` und `tru_guess`
2. Gib den aktuellen Spielstand aus.
3. Durchlaufe folgende Schritte solange, bis die Anzahl der Rateversuche überschritten ist oder bis das Lösungswort erraten wurde:
    1. Nimm einen geratenen Buchstaben entgegen. Hier muss auch überprüft werden, ob die Eingabe *gültig* ist. Überlegen Sie sich dazu, was die Funktion ` take_guess` zurückliefert, wenn kein gültiger Buchstabe eingegeben wird.
    2. Wenn der Buchstabe richtig geraten ist, hänge den Buchstaben an die Liste `tru_guess` an und stelle den neuen Spielsatnd dar.
    3. Wenn der Buchstabe falsch geraten ist if `guess` is wrong, erhöhe den den Zähler `attempt` and Stelle die neue Spielsituation dar.
4. Wenn die Anzahl der maximalen Versuche überschritten wurde, informiere die Spielerin oder den Spieler über ihre oder seine Niderlage.
4. Falls die Anzahl der maximalen Versuche nicht überschritten wurde, beglückwünsche die Spielerin oder den Spieler zu ihrem oder seinem Sieg.

<!-- grade_id: cell-a78967b438aa68ca, points: 0, grade: false, solution: true -->
```python
def play():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c0a584d26fd9bcb0, points: 2, grade: true, solution: false -->
```python
# Test Cell
#----------
@patch('sys.stdout', new_callable=StringIO)
@patch('__main__.take_guess')
@patch('__main__.random_word', return_value = 'OOOO')
def test_play(guess,w,g,s):
    g.return_value = guess
    play()
    return s.getvalue()
#----------
# Winning case
win = test_play('O')

for n in [dis_hangman(0), dis_word('OOOO'), dis_word('OOOO', 'O')]:
    assert n in win, 'Game elements are not displayed as discribed!'

assert win.count(dis_hangman(0)) == 2, 'State of the game should be displayed at the biginning of the game and at winning!'
#----------
# Losing case
lose = test_play('U')

for i in range(6):
    assert dis_hangman(i) in lose, 'All attempts must be exhausted, before the player loses!'

assert lose.count(dis_word('OOOO')) == 6, 'At each stage, the empty underscores must be displayed!'
assert lose.count(dis_word('OOOO', 'O')) == 1, 'The correct must be displayed once to the user, after they lose!'
```

## Fast fertig!

Sie können nun Ihr Spiel ausprobieren!

```python
# Note: rerun all cell if something unexpected occurred!
    
play()
```