<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

# Funkionen

Funktionen in Python werden über das Schlüsselwort `def` definiert. Die Syntax einer Funktions-Definition sieht folgendermaßen aus:

```python
def myfunc(arg1, arg2,... argN):  
  '''Dokumentation'''  

  #Programmcode  

  return <Rückgabewert>  
```

Hier wird die Funktion "myfunc" definiert, welche mit den Parametern "arg1,arg2,....argN" aufgerufen werden kann. Im Funktionsrumpf kann vor dem eigentlichen Programmcode noch eine Dokumentation zu der Funktion angegeben werden. (Siehe auch: Kapitel 36 _Dokumentation_ in: Ernesti und Kaiser, "Python 3: Das umfassende Handbuch", Rheinwerk 2018)

Die Dokumentation sowie der Rückgabewert sind optional. Eine Funktion muss aber mindestends eine Anweisung enthalten. Deshalb ist die folgende Definition ungültig:

```python
def leer():
      
```

Falls Sie eine Funktion definieren, aber nicht ausimplementieren wollen, können Sie die `pass`-Anweisung benutzen. Sie wird eingesetzt, wenn die Syntax eine Anweisung verlangt, das Programm jedoch nichts tun soll.

```python
def fastleer():
    pass
```

## Schreiben einer Funktion

Hier ist noch einmal die Beispiel-Funktion aus Kapitel 19.1 des Lehrbuchs.

```python
def fak(zahl):
    '''Berechnet die Fakultät einer Zahl'''
    ergebnis = 1
    for i in range(2, zahl+1):
        ergebnis *= i
    print(ergebnis)

#Ein einfacher Test:
fak(10)
```

Diese Funktion ist allerdings nur bedingt nützlich, da sie das Resultat der Berechnung nur auf die Standardausgabe schreibt, aber keinen Rückgabewert definiert. Daher ist es sinnvoll, die Funktion so zu umzuschreiben:


```python
def fak(zahl):
    '''Berechnet die Fakultät einer Zahl'''
    ergebnis = 1
    for i in range(2, zahl+1):
        ergebnis *= i
    return ergebnis

#Ein einfacher Test:
fak(10)
```

Natürlich können Sie die Fakultät-Funktion auch rekursiv definieren:

```python
def fak(n):
    '''Berechnet die Fakultät der Zahl n'''
    if n > 0:
        return fak(n - 1) * n
    else:
        return 1
    
print(fak(10))
```

Da wir einen *Docstring* angegeben haben, können wir Informationen zur Funktion über das Hilfesystem abfragen:

```python
help(fak)
```

Auf den Docstring kann man auch innerhalb eines Programms zugreifen:

```python
print(fak.__doc__)
```

**Aufgabe 1**

**Schreiben Sie eine Funktion `ggt(a,b)`, die den größten gemeinsamen Teiler (ggT) zweier Zahlen berechnet.**

**Hinweise:** Eine effiziente Möglichkeit, den ggT zu berechnen ist der euklidische Algorithmus. Das Verfahren funktioniert folgendermaßen:
1. Spezialfälle abfangen: Wenn (a==0) gilt, ist das Ergebnis der Funktion `b`; wenn (b==0) gilt, ist das Ergebnis `a`.  
2. Solange `b!=0` ("b ungleich 0") gilt, wiederhole folgende Schritte:  
    1. Berechne r als den ganzzahligen Rest der Division $\frac{a}{b}$ ("r = a modulo b")  
    2. Setze a=b und b=r
3. Gebe a als Ergebnis zurück


Weitere Informationen finden Sie z.B. auf [Wikipedia](https://de.wikipedia.org/wiki/Gr%C3%B6%C3%9Fter_gemeinsamer_Teiler)

<!-- grade_id: cell-0552cdc5f259730f, points: 0, grade: false, solution: true -->
```python
def ggt(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e21595e4a64976be, points: 1, grade: true, solution: false -->
```python
assert ggt(22,0) == ggt(0, 22) == 22, 'Catch special cases!'
test_cases = [(123, 321, 3),(545,654, 109),(459,666, 9)]
assert all([ggt(a,b)==c for a, b, c in test_cases]), ''
```

Überprüfen Sie die Funktion anhand von Beispielen.

```python
print("ggT von 168 und 546 ist %d (richtig ist 42)" % ggt(168,546))
```

## Funktionsparameter

Sie haben nun eine `ggt` Funktion, die den größten gemeinsamen Teiler zweier Zahlen berechnet. Falls Sie den ggT von mehr als 2 Zahlen berechnen wollen, können Sie die Funktion `ggt` mehrfach benutzen.  
Nehmen wir an Sie wollen den ggt der Zahlen a, b, und c berechnen. Es gilt $ggT(a,b,c)=ggT(ggT(a,b),c)$  
Da für die ggT Funktion das Assoziativgesetz gilt, ist es unerheblich, in welcher "Reihenfolge" Sie die Funktion aufrufen. Es gilt also $ggT(ggT(a,b),c)=ggT(a,ggT(b,c)$.

Um die ggT Funktion nun allgemeiner, für mehrere Parameter zu definieren, benennen wir zunächst die ursprüngliche Funktion `ggt` in `ggt2` um:

```python
ggt2=ggt
```

Das funktioniert, weil der Funktionsname nur ein Zeiger auf das eigentliche Funktionsobjekt ist. Sie können nun die Funktion über beide "Namen" aufrufen:

```python
print(ggt(168,546))
print(ggt2(168,546))
```

Der "alte" Name kann weiter bestehen, Sie können ihn aber auch explizit löschen:

```python
del ggt
print(ggt(168,546))
print(ggt2(168,546))
```

**Aufgabe 2**

**Schreiben Sie eine Funktion `ggt3(a,b,c)`, die den ggT dreier Zahlen a, b und c berechnet. Die Funktion soll auch ein korrektes Ergebnis zurückgeben, wenn nur 2 der 3 Parameter beim Funktionsaufruf gesetzt werden.**

<!-- grade_id: cell-7fe9b0fc2f92e21d, points: 0, grade: false, solution: true -->
```python
def ggt3(a=0, b=0, c=0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f8495582928dda65, points: 1, grade: true, solution: false -->
```python
try:
    ggt3()
except NameError:
    print('ggt function has been deleted, use ggt2 to refer to the function object of ggt.')
    raise

assert ggt3(545,654) == ggt3(b=545,c=654) == ggt3(a=545,c=654), 'Incorrect!'
```

Überprüfen Sie ihre Funktion anhand einiger Beispiele:

```python
ggt3(168,546)
```

```python
ggt3(b=168,c=546)
```

Die ggT-Funktion um einen Parameter zu erweitern ist nicht gerade sinnvoll. Besser wäre es, wenn man `ggt()` mit einer beliebigen Anzahl von Parametern aufrufen könnte.  

**Aufgabe 3**

**Schreiben Sie eine Funktion `ggt`, die den ggT von 2 und mehr Zahlen berechnet.**

**Hinweise:** Verwenden Sie variable (Positions-) Parameter mit dem `*<parameter-tupel>` Konstrukt.

<!-- grade_id: cell-3b8fec3cff001ac9, points: 0, grade: false, solution: true -->
```python
def ggt(a, b, *p, **kwp):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0ccc5e186ba71ef7, points: 1, grade: true, solution: false -->
```python
assert ggt(120, 540) == 60, 'implement the case of two parameters.'
assert ggt(120, 540, 1200, 3420, 5620) == 20, 'implement multiple parameters using p list.'
assert ggt(120, 540, 1200, 3420, c=5620) == 20, 'catch keyworded parameters using kwp dict.'
```

```python
tupel = ()
if tupel:
    print("Nicht leer")
```

Testen Sie ihre Implementierung:

```python
ggt(16, 20, 4, letzter=2)
```

Sie können sich auch zu Testzwecken eine Eingabe für ihre Funktion generieren. Im nächsten Code Abschnitt wird eine Liste von 4er Potenzen erzeugt. Dieses Beispiel zeigt im Übrigen auch eine sinnvollere Anwendung von benannten Parametern: Die `print` Funktion gibt im Normalfall ihren ersten Parameter auf der Standardausgabe aus und schließt die Ausgabe mit einem Zeilenumbruch ab. Wenn Sie statt des Umbruchs eine andere Zeichenkette verwenden wollen, so können Sie den `end` Parameter überschreiben.

```python
for z in range(1,10): print(4**z, end=", ")
print(4**10)
```

```python
ggt(4, 16, 64, 256, 1024, 4096, 16384, 65536, 262144)
```

Im vorherigen Schritt haben wir die eine "Liste" von Argumenten generiert. Diese Liste ist aber nur eine Reihe von ausgegebenen Werten, kein Listen-Objekt in Python.

**Aufgabe 4**

**Wandeln Sie den Test-Code für die `ggt` Funktion so um, dass sie ein Tupel-Objekt `t` and den Funktionsaufruf übergeben. Testen Sie ihre Implementierung**


<!-- grade_id: cell-88af1eb82eb4fb43, points: 0, grade: false, solution: true -->
```python
t = None
# YOUR CODE HERE
raise NotImplementedError()
ggt(4,16,*t)
```

<!-- grade_id: cell-78a0d93dfb3a88b0, points: 1, grade: true, solution: false -->
```python
assert type(t) is tuple, 't should be a tuple!'
assert len(t)>5, 'generate longer sequence!'
assert ggt(4,16,*t) == 4, 'the generated list should be powers of 4'
```

Sie können zusätzlich eine variable Liste von benannten Objekten as Funktionsargument übergeben.
Der Parameter in fer Funktion wird dann mit zwei Sternen `**` eingeleitet.

```python
def myfunct(**kwargs):
    print(list(kwargs.keys()))
    print(list(kwargs.values()))

myfunct(A=1, B='Zwei', C=3)
```

Dabei können Sie auch ein Dictionary übergeben, bei dem die Keys ausschließlich aus Stings bestehen

```python
dictionary = {'Erster':1, 'Zweiter':'2ter', 'Dritter':3.0}
myfunct(**dictionary)

```

Wenn Sie eine variable Liste unbenannter Objekte sowie eine variable Liste benannter Ojekte übergeben wollen, müssen alle unbenanten Objekte **vor** den benannten stehen.

```python
def myfunct(*args, **kwargs):
    print(list(args))
    print(list(kwargs.keys()))
    print(list(kwargs.values()))

myfunct(1, 'B', A=3, B='Vier', C=5)

```

## Namensräume

Beim Aufruf einer Funktion betreten Sie einen neuen Namensraum. Das ist wichtig, damit nicht alle Variablen in Ihrem Programm eindeutig definiert sein müssen.

Wenn Sie also in einer Funktion ein Objekt definieren, dessen Name bereits ausserhalb der Funktion benutzt wurde, hat der *lokale Kontext* Priorität.

```python
a = "a aus main"
def f(a):
    a = "a aus f"
    print(a)

f(a)
a
```

Sie haben allerdings auch innerhalb einer Funktion Zugriff auf den übergeordneten Kontext.

```python
def f():
    print(a)

a = "a aus main"
f()
```

Was der *übergeordnete Kontext* ist, hängt davon ab, wo die Funktion *definiert* wurde; nicht, in welchem Kontext sie aufgerufen wird.

```python
a = "a aus main"

def g():
    print(a)
    
def f():
    a = "a aus f"
    g()

f()
```

**Aufgabe 5**

**Ändern sie die Funktion f so um, dass das Resultat des Aufrufs `f()` "a aus f" ist.**


<!-- grade_id: cell-afc954103555b151, points: 0, grade: false, solution: true -->
```python
%%capture screen --no-stderr --no-display
a = "a aus main"

  
def f():
    # YOUR CODE HERE
    raise NotImplementedError()
f()
```

<!-- grade_id: cell-823f4ddc902e7b21, points: 1, grade: true, solution: false -->
```python
assert 'a aus f' in screen.stdout, 'make sure g is defined in proper domain!'
```

Sobald Sie aus einer Funktion heraus schreibend auf eine globale Variable zugreifen wollen, wird eine lokale Kopie der Variablen erzeugt. Der Wert der Variablen im globalen Kontext bleibt erhalten. Um auch verändernd auf den globalen Namensraum zugreifen zu können, gibt es das Schlüsselwort `global`.

**Aufgabe 6**

**Ändern Sie die Funktion f so um, dass das Resultat des Aufrufs `print(a)` im folgenden Code Abschnitt "a aus f" ist.**

```python
def f():
    a = "a aus f"

a = "a aus main"
f()
print(a)
```

<!-- grade_id: cell-c549f3ae486c687e, points: 0, grade: false, solution: true -->
```python
def f():
    # YOUR CODE HERE
    raise NotImplementedError()

a = "a aus main"
f()
print(a)    
```

<!-- grade_id: cell-4db594722535bf27, points: 1, grade: true, solution: false -->
```python
assert a == 'a aus f', 'Use `global a` to modify a outside the scope of f!'
```

Mit `global` greifen Sie immer auf den globalen Namensraum des Moduls zu. Es kann aber Fälle geben, in denen Sie nicht auf den globalen, sondern den nächst höheren Kontext zugreifen wollen. Dies kann über das Schlüsselwort `nonlocal` erreicht werden. Wenn Sie im folgenden Code Abschnitt `global` durch `nonlocal` ersetzen, erreichen Sie das gewünschte Verhalten.

```python
def f():
    a = "Bitte überschreiben!"
    def g():
        global a
        a = "Neues a"
    print("Altes a in f: ", a)
    g()
    print("Neues a in f: ", a)


a = "Bitte nicht überschreiben!"
f()
print(a)
```

## Anonyme Funktionen

In Python (wie auch in vielen anderen Programmiersprachen) ist es möglich, (Zeiger auf) Funktionen als Argumente an andere Funktionen zu übergeben. Normalerweise muss hierzu die zu übergebende Funktion zuvor definiert werden. Bei Funktionen mit nur sehr wenigen Anweisungen kann dies zu unübersichtlichen Code führen.  
Anonyme Funktionen (oder auch *Lambda Funktionen*) bieten eine Möglichkeit, eine Funktion direkt in einem anderen Ausdruck anzugeben, ohne die Funktion selbst zuvor zu definieren. Eine Summenfunktion kann etwa mit dem Ausdruck `lambda x, y: x+y` beschrieben werden, besitzt aber in dieser Form keinen Namen.  

Es ist aber durchaus möglich, einer Lambda Funktion eine Variable zuzuweisen, über dessen Namen die Funktion im folgenden aufgerufen werden kann:

```python
s = lambda x, y: x+y
s(4,5)
```

Eine "generische Funktion", die eine beliebige Funktion `f` auf 2 Argumente `a` und `b` anwendet und das Resultat zurückgibt, kann folgendermaßen beschrieben werden:

```python
def wasauchimmer(a,b,f):
    return f(a,b)
```

Mittels anonymer Funktionen kann diese "generische Funktion" nun direkt benutzt werden, ohne die jeweiligen Funktionen `f` vorher zu definieren: 

```python
wasauchimmer(4, 5, lambda x, y: x*y)
```

**Aufgabe 7**

**Implementieren Sie die Funktion `mymap` im folgenden Code Abschnitt soweit aus, dass die Funktion `f` auf alle Elemente der Liste `l` angewendet wird. Die so berechneten Elemente sollen in einer neuen Liste zusammengefasst werden, welche als Resultat der Funktion zurückgegeben wird. Testen Sie die Funktion mit 2 verschiedenen Lambda Funktionen ($x\mapsto{}x+1$ und $x\mapsto{}x^2$).**

<!-- grade_id: cell-ed28b3262fff3d34, points: 0, grade: false, solution: true -->
```python
def mymap(l, f):
    # YOUR CODE HERE
    raise NotImplementedError()

l = []
for z in range(10): l.append(z)
print(l)
print(mymap(l, lambda x: x+1))
print(mymap(l, lambda x: x*x))
```

<!-- grade_id: cell-adb875fe5cdcfdcf, points: 1, grade: true, solution: false -->
```python
assert type(mymap([0],lambda x:x)) is list, 'output should be a list with output of lambda on each element of l!'
assert len(mymap(l,lambda x:x)) == len(l), 'l and mymap output should be the same length!'
assert mymap(l, lambda x: x+1) == [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
assert mymap(l, lambda x: x*x) == [0, 1, 4, 9, 16, 25, 36, 49, 64, 81]
```