<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

### Aufgabe: Caesar-Verschlüsselung
In dieser Aufgabe soll eine Klasse zur Verschlüsselung von Textnachrichten ertellt werden.

Wir verwenden dazu ein sehr einfaches Verschlüsselungsverfahren, nämlich die [Caesar-Verschlüsselung](https://de.wikipedia.org/wiki/Caesar-Verschl%C3%BCsselung).
Bei der Caesar-Verschlüsselung wird jedes Zeichen der Textnachricht einzeln verschlüsselt, indem es einfach um einen festen Wert $k$ im Alphabet *verschoben* wird.
Wenn das Alphabet nur aus den Kleinbuchstaben besteht, das zu verschlüsselnde Zeichen ein `a` ist und die *Verschiebungskonstante* $k=5$ gewählt ist, so ist das Ergebnis der Verschlüsselung ein `f`. Um das Zeichen wieder zu entschtschlüsseln, müssen Sie vom Zeichenwert die Konstante $5$ abziehen, um wieder das `a` zu erhalten.

Natürlich kann es passieren, dass Sie bei großen Werten von $k$ oder bei Buchstaben, die recht weit hinten im Alphabet stehen, über den *Rand* des Alphabets heraus kommen. Gleiches gilt anders herum auch für die Entschlüsselung.
Daher müssen Sie bei jedem *Verschieben* eines Buchstabens die Modulo-Operation anwenden. Der Modulo-Divisor muss dabei genauso groß sein, wie die Anzahl der Zeichen im Alphabet.

*Hinweis:* Mit *Alphabet* ist hier die Menge, bzw. in unserem Fall die Liste der betrachteten Zeichen gemeint. Darin enthalten sind also alle *druckbaren* Zeichen. Am einfachsten erhalten Sie diese Liste von Zeichen über das `string` Modul. Dass Attribut `string.printable` liefert Ihnen einen String, der aus allen (standardmäßig) druckbaren Zeichen besteht. In der Mathematik bezeichnet man ein Alphabet häufig mit dem griechischen Buchstaben $\Sigma$ (*Sigma*).

Die Formeln zur Ver- und Entschlüsselung lauten also:

$${\text{encrypt}}_{K}(P)=(P+K)\ \mathrm{mod}\ {N} \ \text{ und }$$

$${\text{decrypt}}_{K}(C)=(C-K)\ \mathrm{mod}\ {N} \ \text{ mit }C, K, P \in \Sigma \text{ und } N=|\Sigma|$$

**Aufgabenbeschreibung:**

1. Erstellen Sie eine Klasse `Cipher`, die als Klassenattribut das Alphabet enthält.
1. Scheiben Sie einen Konstruktor, mit dem Sie ein `Cypher` Objekt erstellen können.
   - Der Konstruktor soll als Parameter einen String als Passwort übergeben bekommen. Es soll später nur mit diesem Passwort möglich sein, eine verschlüsselte Nachricht wieder zu entschlüsseln. Falls im Konstruktor kein Passwort angegeben wird, soll das default Passwort `0000` verwendet werden.
   - Der Konstruktor soll eine zufällige Verschiebungskonstante wählen
   - Der Konstruktor soll außerdem einen Zähler mit $0$ initialisieren. Mit diesem Zählen sollen später fehlerhafte mitgezählt werden.
   - Alle drei Attribute sollen als *private* Instanzattribute angelegt werden, also von *außerhalb* nicht veränderbar sein.
1. Schreiben Sie eine Methode `encrypt`, mit der eine Textnachricht verschlüsselt werden kann. Der Text soll als String-Parameter an die methode übergeben werden. Der String wird dann als Verschlüsselte Zeichenkette unter dem Attribut `encrypted` abgelegt. Dieses Attribut muss nicht privat oder geschützt sein.
1. Schreiben Sie eine Methode `decrypt`, die als Resultat den entschlüsselten Text zurückliefert.
   - Die Methode soll als Parameter das Passwort übergeben bekommen. Nur wenn das Passwort dem gespeicherten Passwort im `Cipher`-Objekt ensricht, soll der Text entschlüsselt werden.
   - Wird ein falsches Passwort übergeben, soll eine Fehlermeldung ausgegeben werden und der interne Zähler um eins erhöht werden.
   - Ist das Passwort zum vierten Mal hintereinander (!) falsch eingegeben worden, soll die verschlüsselte Nachricht gelöscht werden.
1. Wenn ein `Cipher` Objekt als String interpretiert wird (z.B. durch die `print`-Funktion), soll der verschlüsselte Text zurückgeliefert werden.

<!-- grade_id: cell-c3b1669bd31fdede, points: 0, grade: false, solution: true -->
```python
import random
import string

class Cipher:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cd73e516c0641e00, points: 0, grade: true, solution: false -->
```python
secret = Cipher('Skriptsprachen')
secret.encrypt('Viel Spass!')
assert secret.decrypt('Skriptsprachen')=="Viel Spass!"
```

**Erweiterung I:**

Es Soll eine Klasse `ExtCypher` (für *extendable Cypher*) erstellt werden, die von der Klasse `Cipher` erbt.
Bei `ExtCypher` Objekten soll der verschlüsselte Text erweiterbar sein.
Wie bei Strings soll auch bei dieser Klasse der `+=`-Operator verwendet werden können.

Mit `secret+=" Neuer Text"` soll der verschlüsselte Text im `ExtCypher` Objekte `secret` um den verschlüsselten Text " neuer Text" erweitert werden.

*Hinweis:* Da wir in der Klasse `Cypher` einige private Attribute verwendet haben, funktioniert das Erweitern des Strings nicht ohne Weiteres. Welches Attribut muss in der Sichtbarkeit verändert werden, damit das Anhängen funktioniert? Ändern Sie die Klasse `Cipher` so geringfügig wie möglich ab.

<!-- grade_id: cell-05a8d147952445b6, points: 0, grade: false, solution: true -->
```python
class ExtCipher(Cipher):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8ecb4f4380ba6c94, points: 0, grade: true, solution: false -->
```python
secret = ExtCipher('Skriptsprachen')
secret.encrypt('Viel')
assert secret.decrypt('Skriptsprachen')=="Viel"
secret+=" Spass!"
assert secret.decrypt('Skriptsprachen')=="Viel Spass!"
```

**Erweiterung II:**

Es Soll eine Klasse `ICypher` (für *interactive Cypher*) erstellt werden, die von der Klasse `Cipher` erbt.
Statt im Klartext soll nun die Eingabe des Passworts immer interaktiv erfolgen.

Normalerweise können Sie Benutzereingaben mit der Funktion `input()` interaktiv abfragen.
Dabei erscheint allerdings der eingegebene Text im Klartext.
Über das Modul `getpass` und die darin enthaltene Funktion gleichen Namens, können Sie maskierte Eingaben tätigen.
Verwenden Sie also die Funktion `getpass()` um die Passworte abzufragen.

Schreiben Sie möglichst wenig zusätzlichen Code und übernehmen Sie möglichst Methoden und Attribute der Elternklasse.

<!-- grade_id: cell-b8845118e896ad60, points: 0, grade: false, solution: true -->
```python
from getpass import getpass

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-745c246fa8038ee3, points: 0, grade: true, solution: false -->
```python
def _getpass():
    return "Test"

orig = getpass
getpass = _getpass
secret = ICipher()
secret.encrypt('Hallo Welt!')
print(secret.encrypted)
print(secret.decrypt())
assert secret.decrypt()=="Hallo Welt!"
getpass = orig
```

```python

```