<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

## Sortieren

Eines der klassischen Probleme der Informatik ist das Sortieren von Daten. Es existierende
dutzende Sortieralgorithmen mit unterschiedlichen Eigenschaften. Ein recht einfaches Sortierverfahren
ist *Selection Sort*, ein in der Praxis (und bei größeren Datenmengen) sehr schnelles Verfahren ist *Quicksort*

In diesem Aufgabenblatt wollen wir die beiden genannten Sortierverfahren mit Python implementieren.


## Vorbereitung

Um zu Sortieren benötigt man natürlich einen Datensatz, den es zu sortieren gilt.
Häufig verwendet man dazu Arrays (oder Listen) die mit zufälligen Werten gefüllt sind.

Um Zufallszahlen mit Python zu generieren, kann man das Paket `random` verwenden.
Die folgende Zelle erzeugt zunächst ein Array mit 1.000 zufälligen Zahlen und stellt diese mit Hilfe von `matplotlib` dar.
(Diese Bibliothek haben wir schon kurz beim letzten Termin gesehen)

```python
# Wir benötigen Zufallszahlen und eine Bibliothek zum plotten
%matplotlib inline
import random
import matplotlib.pyplot as plt

N = 1000

# Eintausend Zufallszahlen zwischen 0 und 1
unsorted = [ random.random() for i in range(N)]

# Die folgenden Zeilen zeigen die Zahlen als Scatter-Plot an
plt.scatter(range(N), unsorted)
plt.show()
```

## Selectionsort


*Selectionsort* funktioniert so, dass es eine Liste in zwei nebeneinanderliegende Teillisten aufteilt.
Die linke Teilliste ist sortiert und hat am Anfang die Größe 0.
Die rechte Teilliste ist unsortiert und hat am Anfang die Länge $N$, wobei $N$ die Anzahl der Elemente in der Liste ist.
In jedem Schritt sucht *Selectionsort* das kleinste Element und dessen Position in der noch unsortierten Liste.
Dazu müssen alle Elemente in der rechten Liste einmal angeschaut werden.
Hat man das kleinste Element der unsortierten Liste gefunden, so vertauscht man es mit dem ersten Element in der unsortierten Liste.
Danach ist die sortierte Liste um ein Element größer geworden und die unsortierten Liste um ein Element kleiner.
Im nächsten Schritt sucht man dann das kleinste Element in der verblieben unsortierten Liste.
Diese Abfolge wird solange wiederholt, bis die unsortierte Liste nur noch ein Element groß ist.
Danach ist die Liste vollständig sortiert.

Schreiben Sie eine Funktion `selectionsort(l)` zum Sortieren der liste `l`.


<!-- grade_id: cell-75980d98c92a4e66, points: 0, grade: false, solution: true -->
```python
def selectionsort(l):
    """ Sortiert die Liste l ab dem Element p mit Selectionsort """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a9efe24d56f32c95, points: 1, grade: true, solution: false -->
```python
l = [9,1,4,8,5,7,2,6,3,0]
print("Vorher: ", l)
selectionsort(l)
assert l == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
print("Nachher:", l)
```

## Quicksort

Der *Quicksort* Algorithmus deutlich aufwendiger zu implementieren.
Dafür ist sein durchschnittliches asymptotisches Laufzeitverhalten mit ${\mathcal {O}}(N\cdot \log(N))$ deutlich besser als das einfacher Sortieralgorithmen, wie etwa *Selcetion Sort* (${\mathcal {O}}(N^2)$).

Schreiben Sie eine Funktion `quicksort(lst)`, die eine übergebene Liste sortiert.
VErwenden Sie **nicht** die eingebauten Sortierfunktionen von Python, sondern implementieren Sie den Algorithmus selbst.

Als Hilfestellung kann Ihnen vielleicht die folgende *IDEA-Bauanleitung* dienen.
Diese unkonventionelle Beschreibung des Algorithmus stammt von [Sándor P. Fekete](https://www.ibr.cs.tu-bs.de/users/fekete/) von der TU Braunschweig.
Auch wenn die Abbildung auf den ersten Blich amüsant wirkt, beschreibt sie doch recht gut, wie Quicksort funktioniert.

Für eine textuelle Beschreibung (mit Pseudocode) können Sie auch [Wikipedia](https://de.wikipedia.org/wiki/Quicksort) zurate ziehen.

![](quick-sort.png)



<!-- grade_id: cell-016b01013cb5dded, points: 0, grade: false, solution: true -->
```python
def quicksort(l, links=0, rechts=-1, visual=False):
    """Diese Funktion sortiert die übergebene Liste ('in-place sort')
    
    Argumente:
    l      -- zu sortierende Liste
    links  -- Startindex des zu sortierenden Bereich (inklusiv)
    rechts -- Endindex des zu sortierenden Bereichs (exklusiv)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4d336aa434bdc1fb, points: 1, grade: true, solution: false -->
```python
N = 1000
data = [ random.random() for i in range(N)]
quicksort(data)
for i in range(len(data)-1):
    assert data[i] <= data[i+1]
```

## Ergebnis überprüfen
Der folgende Code plotted das Ergebnis und testet, ob die Liste tatsächlich sortiert ist.

```python
N = 1000
data = [ random.random() for i in range(N)]
quicksort(data)

# Die folgenden Zeilen zeigen die Zahlen als Scatter-Plot an
plt.scatter(range(len(data)), data)
plt.show()
```