<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

## Buchstabenrätsel

In dieser Aufgabe geht es darum, herauszufinden, ob ein Wort in einem Buchstabenquadrat gefunden werden kann.
Ein Buchstabenquadrat ist eine quadratische Matrix, deren Werte aus einzelnen Buchstaben bestehen.
(Die Idee dieser Aufgabe basiert auf Übungsaufgaben von Prof. Mehner und Prof. Gawron)

![raetsel.png](raetsel.png)

Um ein Wort zu finden, können Sie folgendermaßen vorgehen:
1. Für den ersten Buchstaben können Sie bei jedem beliebigen Buchstaben des Quadrat starten.
2. Wenn der erste Buchstabe des Wortes dem Buchstaben im Quadrat entspricht, können Sie von dieser Position aus weitersuchen. Dabei dürfen Sie nur zu den (maximal) acht Nachbarfeldern des gefundenen Buchstaben weiterlaufen.

Wie können Sie herausfinden, ob ein gesuchtes Wort im Buchstabenquadrat vorkommt?

## Hinweise

In dieser Aufgabe können **List Comprehensions** sehr hilfreich sein.
Mit List Comprehensions kann man recht kurz und elegant neue Listen-Objekte generieren.
So können wir z.B. eine 2-dimensionale Liste mit Werten füllen.
Mathematisch gesehen, entspricht dies einer Matrix.

```python
k = 4
A = [ [x for x in range(k*y,k*y+k)] for y in range(k) ]
A
```

In einer zweiten Liste können wir nun alle Index-Paare der Matrix erzeugen. Ein Index-Paar ist dabei ein Tupel `(i,j)`.

```python
idx = [ (i,j) for i in range(k) for j in range(k) ]
idx
```

Wenn wir nur an den Indizes der Diagonalen interessiert sind, können wir der *List Comprehension* eine Bedingung hinzufügen.

```python
idx = [ (i,j) for i in range(k) for j in range(k) if i==j]
idx
```

Nun können wir die Index-Tupel verwenden um damit auf die Werte der Matrix zuzugreifen:

```python
diag_values = [ A[i][j] for (i,j) in idx]
diag_values
```

## Implementierung

Beginnen wir zunächst mit der Definition der Matrix:

```python
A = [ list("SRGD"), list("OFAE"), list("NUBL"), list("TSPE") ]
A
```

Output:
```
[['S', 'R', 'G', 'D'],
 ['O', 'F', 'A', 'E'],
 ['N', 'U', 'B', 'L'],
 ['T', 'S', 'P', 'E']]
```

## Mögliche Felder generieren

Ein Feld im Buchstabenquadrat wird eindeutig als Tupel `(i, j)` dargestellt.
Schreiben Sie zwei Funktionen die jeweils eine Liste von Tupeln zurückliefern.
Die Funktion `startfelder(A)` gibt dabei alle Felder der Matrix `A` zurück.
Die Funktion `nachbarn(A,i,j)`, gibt alle möglichen Nachbarfelder des Feldes `(i, j)` von `A` zurück. In der *Mitte* des Buchstabenquadrats sind das acht Felder, am Rand sind es dementsprechend weniger Felder.

<!-- grade_id: cell-44bd158093f8fe3d, points: 0, grade: false, solution: true -->
```python
def startfelder(A):
    """ Gibt eine Liste aller Startfelder (Tupel) zurück """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7b6e3c58ccf274fd, points: 1, grade: true, solution: false -->
```python
A_ = [ list("HAL"), list("LOW"), list("ELT") ]
assert startfelder(A_) == [(0,0),(0,1),(0,2),(1,0),(1,1),(1,2),(2,0),(2,1),(2,2)]
```

<!-- grade_id: cell-d3882e3f11bdbf5c, points: 0, grade: false, solution: true -->
```python
 def nachbarn(A,i,j):
    """ Gibt eine Liste aller Nachbarfelder zurück """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e436fa7f73634464, points: 1, grade: true, solution: false -->
```python
A_ = [ list("HAL"), list("LOW"), list("ELT") ]
assert nachbarn(A_,0,1) == [(0,0),(0,2),(1,0),(1,1),(1,2)]
```

## Suchfunktion

Implementieren Sie nun eine Suchfunktion, die prüft, ob ein gegebenes Wort in der Matrix vorkommt.
Es ist dabei hilfreich, sich das Problem der Suche zunächst abstrakt als Baum vorzustellen.
Sie beginnen mit der Suche des ersten Buchstaben des Wortes an der Wurzel.
Von dort aus haben sie $N^2$ Möglichkeiten, den Buchstaben zu finden.
$N$ ist dabei die Kantenlänge des Buchstabenquadrats mit $N^2$ Feldern.

Zu Beginn gehen Sie von der Wurzel aus zu allen Feldern und überprüfen of der Erste Buchstabe passt.
Die Funktion `startfelder` hilft Ihnen dabei die $N^2$ möglichen Felder zu bestimmen.
Das Ergebnis dieses Schritts ist eine Liste gültiger Startfelder.

Vor allen dieser Felder können Sie nun eine rekursive Suche starten, in der Sie versuchen das Suchwort **ohne den ersten Buchstaben** zu finden.
Am einfachsten funktioniert dies, wenn Sie eine separate Funktion für die rekursive Suche implementieren.
Die rekursive Funktion könnte folgendermaßen arbeiten:
1. Finde alle Nachbarfelder des aktuellen Felds
1. Von den Nachbarfeldern, finde alle Felder die zum ersten Buchstaben des Wortes passen
1. Wenn keine Felder gefunden werden, gib 0 zurück
1. Wenn passender Felder gefunden werden, prüfe ob das Suchwort nur ein Zeichen lang ist.
    1. Wenn ja ist das ursprüngliche Suchwort komplett gefunden. Gib die Anzahl der Treffer-Felder als Ergebnis zurück
    1. wenn das aktuelle Suchwort länger als ein Zeichen ist, müssen noch weitere Buchstaben überprüft werden. Rufe für alle Treffer-Felde die Rekursive Suche mit dem aktuellen Suchwort **ohne den ersten Buchstaben** auf.
1. Gib die Anzahl der gefundenen Worte zurück. Ein und dasselbe Wort kann öfter (d.h. über mehrere verschiedene Pfade) gefunden werden.



<!-- grade_id: cell-9958f94baa5137a4, points: 0, grade: false, solution: true -->
```python
def suche_rek(w,A,i,j):
    felder = nachbarn(A,i,j)
     # YOUR CODE HERE
     raise NotImplementedError()
```

<!-- grade_id: cell-2e0571e310a002f7, points: 1, grade: true, solution: false -->
```python
A_ = [ list("HAL"), list("LOW"), list("ELT") ]
assert suche_rek('OL',A_,0,1)==3
```

<!-- grade_id: cell-2e52da07c25da6f3, points: 0, grade: false, solution: true -->
```python
def suche(w,A):
    felder = startfelder(A)
     # YOUR CODE HERE
     raise NotImplementedError()
```

<!-- grade_id: cell-6c0b37e024045d05, points: 1, grade: true, solution: false -->
```python
A_ = [ list("HAL"), list("LOW"), list("ELT") ]
assert suche('AOL',A_)==3
```

```python
# ORGELBAU sollte gefunden werden
suche("ORGELBAU",A)
```

```python
# SONATE sollte nicht gefunden werden
suche("SONATE",A)
```

```python
# LE sollte 2-mal gefunden werden
suche("LE",A)
```

## Erweiterung

1. In der *echten* Version des Buchstabenrätsels dürfen einmal besuchte Felder nicht erneut besucht werden (kreuzungsfreier Weg). Erweitern Sie die Suchfunktion so, dass Felder nur einmalig besucht werden dürfen.
1.  Das Ziel eines Buchstabenrätsels ist ursprünglich nicht, dass Sie ein bestimmtes Wort finden, sondern dass Sie alle möglichen Worte in dem Buchstabenquadrat aufspüren. Allerdings ist unsere Funktion zum Suchen eines Wortes sehr hilfreich, um auch dieses Problem zu lösen. Wir mussen *nur* alle *möglichen* Worte im Buchstabenquadrat suchen. Eine Liste aller deutschen Worte finden Sie in der Datei `german.lst`. Erweitern Sie das Programm so, dass alle Worte im Buchstabenquadrat gefunden werden. Wie viele sind es?