<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Sommersemester 2021
Prof. Dr. Heiner Giefers

# Objektorientierung

Die Objektorientierung ist das zentrale Element in Python. Alles, was Sie mit einem Bezeichner versehen können, ist schlussendlich eine Instanz einer Klasse. Das gilt für elementare Datentypen wie `int` oder `float`, genauso wie für Funktionen und natürlich Klassen selbst. Abhängig von seiner Klasse hat jedes Objekt spezielle Attribute und Methoden.

Klassen können von anderen Klassen erben; das gilt für selbst-definierte wie auch für eingebaute Klassen. Dabei werden die Methoden der "Oberklasse" in der neuen Klasse zugreifbar und somit wird die Funktionalität der (zumeist generelleren) Oberklasse auf die erbende Klasse übertragen.




Ein neue Klasse wird über das Schlüsselwort `class` definiert. 

```python
class A():
    pass

A()
```

```python
print("A ist vom Typ", type(A))
a = A
print("a ist vom Typ", type(a))
b = a()
print("b ist vom Typ", type(b))
```

## Attribute und Methoden
Innerhalb einer Klasse können Sie (wie in einem Modul oder in einer Funktion) Methoden und Attribute definieren. Wenn Sie in der Klasse `A` z.B. ein Attribut `X` anlegen, können Sie für eine Instanz `a` von `A` über `a.X` auf das `X` in dem Objekt (=der Instanz) `a` zugreifen.

```python
class A():  
    X = 100

a = A()
print(a.X)
a.X=42
print(a.X)
b = A()
print(b.X)
a.Y=12334
print(a.Y)
```

Für Methoden funktioniert das ganz ähnlich, allerdings müssen Sie in Python eine besondere Konvention einhalten: **Der erste Parameter einer Methode ist immer eine Referenz auf die Instanz, von der sie aufgerufen wird.** Innerhalb einer Klasse bezeichnet man das aktuelle Objekt mit dem Schlüsselwort `self`.
Der Parameter `self` ist allerdings nur bei der Definition der Methode erforderlich. Wird eine Methode über ein Objekt aufgerufen (wie im Beispiel unten mit `a.f()`) so wird das Objekt implizit an die Funktion übergeben.

```python
class A():  
    X = 100
    def f(self, i):
        X = i

a = A()
print(a.X)
a.f(42)
print(a.X)
```

Im obigen Beispiel sehen Sie, dass das Attribut `X` des Objekts `a` durch die Funktion `f()` nicht überschrieben wird. Das liegt daran, dass das `X` in `f()` ein lokales Attribut innerhalb der Methode ist. Wenn man das Klassen-Attribut `X` ändern möchte, muss man das Attribut über `self.x` referenzieren.

```python
class A():  
    X = 100
    def f(self,i):
        self.X = i

a = A()
print(a.X)
a.f(42)
print(a.X)
```

Eine Methode kann übrigens auch über Ihren Klassennamen aufgerufen werden. In dem Fall muss aber das Objekt als erster Parameter übergeben werden:

```python
b = A()
print(b.X)
A.f(b,123)
print(b.X)
```

## Konstruktor
Klassen haben üblicherweise eine feste Methode, die beim Erstellen einer Instanz initial aufgerufen wird. Diese Methode nennt man auch Konstruktor.
Ein Konstruktor innerhalb einer Klasse hat immer den Namen `__init__`, wird aber ansonsten wie eine "normale" Methode definiert. Der Konstruktor sollte im Normalfall alle Attribute der Klasse initialisieren.
Es gibt keine (echten) Destruktoren, aber die `del`-Operation und die Magic Method `__del__()`

```python
class A():
    def __init__(self,x):
        self.X = x
    def __del__(self):
        print("Zerstöre Instanz von A")

a = A(100)
print(a.X)
del a
a
```

## Vererbung
Ein wichtiges Konzept in der Objektorientierten Programmierung ist die Vererbung. Neue Klassen können von bestehenden Klassen abgeleitet werden und _erben_ dadurch alle Fähigkeiten ihrer Elternklasse (auch: _Basisklasse_).
Im folgenden Beispiel erbt eine Klasse `B` von `A` und ist damit zunächst einmal eine Kopie der Klasse `A`.

```python
class A():
    def __init__(self):
        self.X = 100

class B(A):
    pass

b = B()
b.X
```

Die abgeleitete Klasse kann dann für ihren eigenen Zweck angepasst und erweitert werden. Es können neue Attribute und Methoden eingeführt werden, oder auch bestehende Attribute und Methoden von der Basisklasse überschrieben werden.

Im Normalfall ist es sinnvoll, einer abgeleiteten Klasse einen eigenen Konstruktor zu geben, indem man die `__init__`-Methode der Basisklasse überschreibt. Hierbei ist es üblich, den Konstruktor aus der Basisklasse mittels `super().__init__()` aufzurufen und so die bestehenden Initialisierungsschritte der Mutterklassse zu übernehmen.
Die Funktion `super()` ermittelt dabei die Basisklasse der aktuellen Klasse.

```python
class A():
    def __init__(self):
        self.X = 100

class B(A):
    def __init__(self):
        super().__init__()
        self.Y = 200

b = B()
print(b.X, b.Y)
```

Im Unterschied zu Java (aber ähnlich zu C++) ist in Python **Mehrfachvererbung** möglich. Es können also bei der Klassen-Definition mehrere Basisklassen angegeben werden. Dadurch ergeben sich Fragestellungen, vor allem bezüglich der Eindeutigkeit von Referenzen.

Solange eine Klasse nur von einer anderen Klasse erbt, sind Bezeichner immer eindeutig. Ein Funktion, die von einer Basisklasse bereit gestellt wird, wird in der neuen Klasse entweder übernommen oder überschrieben. Bei Mehrfachvererbung kann es aber dazu kommen, dass Methoden (gleiches gilt für Attribute) gleichen Namens von verschiedenen Basisklassen geerbt werden. In diesem Fall ist nicht mehr eindeutig, welche Funktion in einem Aufruf durch ein Objekt der Unterklasse gemeint ist. Ein Beispiel:

```python
class O():
    def f(self):
        print("f() aus O")
class A(O):
    pass

class B(O):
    def f(self):
        print("f() aus B")

class C(B,A):
    pass

c = C()
c.f()
```

In diesem Beispiel ist es nun nicht mehr so klar ersichtlich, welche `f`-Methode am Ende aufgerufen wird. Daher ist die Reihenfolge, in der nach möglichen Referenzen für ein Methodenaufruf in den Oberklassen einer Klasse gesucht wird,  durch die sogenannte _Method Resolution Order_ festgelegt.

Man kann die geerbten Methoden auch direkt über Ihren Klassennamen ansprechen. Dies ist z.B. für das Aufrufen mehrerer Konstruktoren von Basisklassen sinnvoll.

```python
class A():
    def __init__(self):
        print("ich bin ein A")
class B():
    def __init__(self):
        print("ich bin ein B")

class C(A,B):
    def __init__(self):
        A.__init__(self)
        print("und")
        B.__init__(self)
c=C()
```

## Zugriffsregeln

Zugriffsmodifikatoren, so wie etwa _public_ und _private_ in Java, gibt es in Python nicht. Sie können (technisch gesehen) jedes Attribut eines Objektes "von Außen" über die Objekt-Referenz zugreifen und auch überschreiben.
Es gibt allerdings einige Konventionen, die Python-Entwickler benutzen, um zu signalisieren, dass Attribute einen "geschützten" Status haben sollen. So ist es üblich, dass alle Attribute, die mit einem `_` (Unterstrich) beginnen,  außerhalb der Klassendefinition nicht direkt zugegriffen werden sollen. Für solche Attribute können ggf. _setter_ und _getter_ Methoden implementiert werden.
Über die eingebaute Funktion `property([fget,fset,fdel,doc])` kann ein Property-Attribut erzeugt werden, über dessen Namen man dann einfachen Zugriff auf die getter und setter Methoden (sowie eine delete Methode und einen _Docstring_ zum Attribut) hat.

```python
class A():
    def __init__(self):
        self._X = 100
        self.__X = 200
        
    def get__X(self):
        print("Get __X")
        return self.__X
    
    def get_X(self):
        print("Get X")
        return self._X
    def set_X(self,x):
        print("Set X")
        self._X = x
    def delX():
        del self._X
    X = property(get_X,set_X,None,'Ein geschütztes Attribut X')

a=A()
a.X = 42
print(a._X)
print(a.X)
a._X = 123
print(a.X)
help(A.X)
```

Im obigen Beispiel wird neben dem Attribut `_X` ein zweites Attribut `__X` angelegt. Sind dem Attributnamen zwei Unterstriche vorangestellt, ändert der Python Interpreter den Namen intern um, so dass er nicht mehr direkt von außerhalb der Klasse zugegriffen werden kann.

```python
print(a.get__X())
a.__X
```

```python
a._A__X
```

## Statische Methoden und Klassenmethoden
Eine Methode aus dem globalen Namensraum kann über die eingebaute Funktion `staticmethod()` als statische Methode in eine Klasse übernommen werden. Sie kann dann über die Klasse selbst oder Objekte der Klasse aufgerufen werden.

```python
class A():
    def f():
        print("Hallo von f()")
    
    def h(self):
        pass
    g = staticmethod(f)

A.g()
a = A()
a.g()
```

Ähnlich zu statischen Methoden verhalten sich Klassenmethoden. Auch sie beziehen sich nur auf die Klasse selbst und nicht auf Instanzen der Klasse. Im Unterschied zu statischen Methoden werden Klassenmethoden im Namensraum der Klasse definiert und erwarten als ersten Parameter (in der Funktionssignatur) eine Referenz auf die Klasse. Klassenmethoden können also nicht alleine stehen, sondern gehören immer zu einer Klasse.

```python
class A():
    def f(c):
        print("Hallo von %s" % c)
    g = classmethod(f)
class B(A):
    pass

A.g()
a = A()
b = B()
a.g()
b.g()
```

Auch Klassenattribute sind bereits in der Klasse definiert und werden bei einer Objekterzeugung in die Instanz übernommen. Auch sie können, ohne dass eine Instanz existiert, über den Namen der Klasse zugegriffen werden.

```python
class A():
    X = 100

class B(A):
    pass

a = A()
b = B()
b.X = 200
print(a.X, B.X, b.X)
```

## Magic Methods
In Python gibt es eine ganze Reihe von speziellen Attributen und Namen mit einer vordefinierten Funktion. Das heißt, der Interpreter sucht bei bestimmten Operationen nach Methoden oder Attributen mit einem festdefinierten Namen. Sie können diese Methoden und Attribute überschreiben, um die vordefinierte Funktionalität zu ersetzen oder zu erweitern.
Die `__init__(self,...)` Funktion ist ein Beispiel für eine Magic Method. Wenn Sie eine neue Instanz einer Klasse erzeugen, sucht der Interpreter nach einer Methode mit genau diesem Namen und führt sie aus.
Es gibt auch Attribute mit vordefinierter Bedeutung. Der _Docstring_ `__doc__` ist ein Beispiel für solch ein Magic Attribute.

```python
class A():
    '''Eine Klasse A'''
    def __init__(self):
        self.X = 100

class B(A):
    pass

a = A()
b = B()
print(b.X)
print(A.__doc__)
B.__name__
```

Besonders hilfreich sind Magic Methods um sehr kompakten Code zu ermöglichen. So können etwa Operatoren wie `+`, `-`, `*`, ... überladen werden, indem man die Methoden `__add__`, `__sub__`, `__mul__`, ... implementiert. Damit können Sie z.B. eine Plus-Operation für Ihre neue Klasse völlig frei definieren. Diese Methode wird immer dann aufgerufen, wenn Sie zwei Objekte der Klasse mit einem `+`-Symbol verbinden.

```python
class A():
    def __init__(self,x):
        self.X = x
    def __add__(self,other):
        neu = A("")
        neu.X = self.X + " und " + other.X
        return neu

a = A("Ein A")
b = A("noch ein A")
(a+b).X
```

## Aufgaben

Für die Aufgaben zum Thema "Objektorientierung" benutzen wir (so wie das Lehrbuch) ein Beispiel zur Kontoverwaltung.
Es sollen mehrere Klassen entwickelt werden, die teilweise voneinander erben.
In der ersten Aufgabe geht es um eine Klasse, die den Wert des Kontostands speichert. Hierfür ist es sinnvoll, sich zunächst ein Phänomen anzuschauen, das in vielen Bereichen der computergestützten Mathematik wichtig ist.

Normalerweise benutzen Computer eine Binärdarstellung um Zahlen zu kodieren. Das gilt für alle Zahlen, egal ob ganzzahlige oder Fließkomma Datentypen. Allerdings treten bei der Darstellung von Fließkommazahlen Rundungsfehler auf, sofern sich der Betrag nicht genau als Summe von 2-er Potenzen mit der Anzahl der verfügbaren Bits darstellen lässt. Je mehr Bits für die Speicherung der Zahl benutzt werden, umso gnauer die Darstellung.
Python benutzt zur internen Darstellung von `float` Zahlen den 64-bit `double` Datentyp, den Sie z.B. aus C++ kennen.
Trotz dieser "doppelten" Präzision, kann es zu signifikanten Berechnungsfehlern kommen:

```python
a = 1.1; b = 1.3
print("%f" % (a-b))
(a-b)
```

Wenn sehr viele Einzeloperationen durchgeführt werden, kann sich diese kleine Ungenauigkeit schnell aufsummieren und so ein erheblicher Fehler entstehen.
Betrachten Sie folgendes Beispiel: Sie haben ein "leeres" Bankkonto und überweisen dorthin zunächst 1,10 EUR und buchen dann direkt wieder einen Euro ab. Wenn Sie sich den Kontostand anzeigen lassen, sehen Sie, dass etwas mehr als die erwarteten 10 Cent auf dem Konto sind.

```python
kontostand = 0.0
kontostand += 1.1
kontostand -= 1.0    
print("%s Euro" % kontostand)
```

Wenn Sie diesen Vorgang sehr oft wiederholen und dann den ursprünglich eingezahlten Betrag komplett abziehen, sehen Sie, dass ein erhebliches Plus auf dem Konto verbleibt.

```python
restbetrag = kontostand
anzahl_transaktionen = 1e20
kontostand = anzahl_transaktionen * restbetrag
summe_komplett = 0.1 * anzahl_transaktionen
kontostand -= summe_komplett
print("%s Euro" % kontostand)
```

Bei einem anderen Datenformat, z.B. 32-bit `float` treten Rundungsfehler dieser Art schon nach sehr viel weniger Einzeloperationen auf.
Es ist daher ratsam, für die Darstellung von Dezimalzahlen ein eigenes Datenformat zu implementieren, das die Vorkomma- und Nachkommastellen einer Dezimalzahl exakt als Integer abspeichert.

**Aufgabe 1**

**Entwickeln Sie eine Klasse `Dezimal` die eine Fließkommazahl als Dezimalzahl darstellt.** 

Sie sollten in Ihrer Klasse die Vorkomma-, wie die Nachkommastelle getrennt als Integer-Zahlen abspeichern. Bei dem Wert der Nachkommastelle ist zu beachten, dass der Wert abhängig von der Anzahl der Nachkommastellen ist. Die Zahl der Nachkommastellen soll mit 6 festgelegt sein. Die Zahl 3,04 hätte also die Attribute `Vorkomma=3` und `Nachkomma=40000`.

Die Klasse `Dezimal` sollte einen Konstruktor besitzen, der als Parameter eine Fließkommadarstellung übergeben bekommt. Darüberhinaus sollte die Klasse eine Methode besitzen, die den gespeicherten Wert als formatierten String  zurück gibt.
Hierzu noch einige Hinweise:

Mit dem Formatierungs-String "%0.6d" können Sie einen Integer-Wert mit mindestens 6 Stellen ausgeben:

```python
gross = 123456789
klein = 12
print("gross = %0.6d und klein = %0.6d" % (gross,klein))
```

Sie sollten beachten, dass `int(x)` eine Fließkommazahl nach dem Komma "abschneidet", also immer abrundet. Bei der Nachkommastelle sollten Sie aber numerisch korrekt Runden um Fehler zu vermeiden.

```python
x = 1.78
print("x = %f als int ist %d und gerundet %d" % (x, int(x), int(round(x))))
```

Übrigens, sehr große Integer Werte sind für Python kein Problem. Falls die Zahlen "zu groß" für einen 32-/64-Bit Integer Datentyp sind, verwendet Python die _bignum_ Bibliothek. 

```python
mybignum = 28972348923472352359237
mybignum*9792
```

<!-- grade_id: cell-33dfe77bd3caa061, points: 0, grade: false, solution: true -->
```python
class Dezimal():
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-334a72d49bd14e3e, points: 1, grade: true, solution: false -->
```python
test_cases = [-12, 0.33,5.0000009, 100]
for c in test_cases:
    t = Dezimal(c)
    assert t.Vorkomma == int(c), 'Vorkomma is not correct!'
    assert t.Nachkomma == int(round((c-int(c))*1000000)), 'Nachkomma is not correct!'

t = Dezimal(-22.0000559)
tt = str(t)
assert tt.count('-') == 1, 'Make sure the sign appears only before the number!'
assert len(tt) == 10, 'Make 6 places after the decemal point.'
```

```python
a = Dezimal(12.3)
print(a)
```

**Aufgabe 2**

**Entwickeln Sie eine Klasse `Kontostand`, die von der Klasse Dezimal erbt. Die Klasse soll einen eigenen Konstruktor haben, der zusätzlich zum Betrag des Kontostands noch ein Attribut `Datum` initialisiert. Sie können das Datum einfachheitshalber als String anlegen.**

<!-- grade_id: cell-668d88cf645c76b8, points: 0, grade: false, solution: true -->
```python
class Kontostand(Dezimal):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0f1640032540cb33, points: 1, grade: true, solution: false -->
```python
try:
    t = Kontostand(22.05)
    t.Nachkomma
except TypeError:
    print('Set up a default date!')
    raise
except AttributeError:
    print('Include the initialization of Dezimal using super().')
    raise
    
t = Kontostand(22.05, '11.12.2019')
assert t.Datum in str(t), 'Include the date in the class string!'
assert str(Dezimal(22.05)) in str(t), 'Include the amount in the class string!'
```

```python
a = Kontostand(12.3)
print(a)
```

**Aufgabe 3**

**Erweitern Sie Ihre Klasse `Dezimal` um eine Methode `plus`, die die Summe zweier Dezimalzahlen berechnet und als neues `Dezimal`-Objekt zurück gibt. Überführen Sie Ihre Methode `plus` in eine "Magic Method", so dass Sie 2 `Dezimal`-Objekte mittels des `+`-Operators addieren können.**

<!-- grade_id: cell-0442a1a52b5af383, points: 0, grade: false, solution: true -->
```python
class Dezimal():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e912e52bc0977b26, points: 1, grade: true, solution: false -->
```python
t0, t1 = Dezimal(11.11), Dezimal(2.9)
assert (t0+t1).Vorkomma == 14, 'Vorkomma is not correct!'
assert (t0+t1).Nachkomma == 10000, 'Nachkomma is not correct!'
assert len(str(t0+t1)) == 9, 'The resulting string should be 9 characters long.'
```

```python
a = Dezimal(12.3)
b = Dezimal(11.6)
print("%s + %s = %s" % (a, b, a+b))
print("%s - %s = %s" % (a, b, a-b))
```

**Aufgabe 4**

**Erweitern Sie Ihre Klasse `Kontostand` um die Methoden `einzahlen` und `auszahlen` mit denen Sie den Kontostand erhöhen oder verringern können.**

<!-- grade_id: cell-9254fdc7699551eb, points: 0, grade: false, solution: true -->
```python
import time
def timestamp():
    return time.strftime("%d.%m.%Y, %H:%M Uhr")

class Kontostand(Dezimal):
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5016897b28bef218, points: 1, grade: true, solution: false -->
```python
t = Kontostand(1.0)
t.einzahlen(10.5)
assert t.Vorkomma+t.Nachkomma/10**6 == 11.5, 'Deposit was not correctly processed!'
t.auszahlen(5.6)
assert t.Vorkomma+t.Nachkomma/10**6 ==5.9, 'Withdrawal was not correctly processed!'
assert t._Datum == timestamp(), 'Update transaction timestamp!'
```

```python
a = Kontostand(0.0, '01.05.2018')
print(a)
a.einzahlen(12.3, datum='11.05.2018')
print(a)
a.auszahlen(8.6)
print(a)
```

**Aufgabe 5**

**Ergänzen Sie die Klasse Kontostand um eine statische Methode, die die Informationen "Dies ist ein EUR Konto" ausgibt.**

<!-- grade_id: cell-aa804db3315618ea, points: 0, grade: false, solution: true -->
```python
# add static method "info()"

class Dezimal():
    # YOUR CODE HERE
    raise NotImplementedError()
    
class Kontostand(Dezimal):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d96190b1d47a0442, points: 1, grade: true, solution: false -->
```python
Kontostand(0.0).info()
Kontostand.info()
```

```python
Kontostand.info()
a = Kontostand(0.0, '01.05.2018')
a.info()
```

#### Aufgabe 6
**Schreiben Sie eine Klasse `Zeitstempel`, die ein Datum (als String) speichert. Die Klasse soll eine Funktion `aktualisieren()` haben, mit der das interne Datum auf die aktuelle Zeit gesetzt wird. Benutzen Sie dazu die folgende Funktion `timestamp()`. Verwenden Sie diese Klasse in der Form, dass Kontostand nun auch von `Zeitstempel` erbt.**

```python
import time
def timestamp():
    return time.strftime("%d.%m.%Y, %H:%M Uhr")
```

<!-- grade_id: cell-8018b28b9a36d756, points: 0, grade: false, solution: true -->
```python
class Zeitstempel():
    # YOUR CODE HERE
    raise NotImplementedError()

class Dezimal():
    # YOUR CODE HERE
    raise NotImplementedError()
    
class Kontostand(Dezimal,Zeitstempel):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-147864f3a3df25f6, points: 1, grade: true, solution: false -->
```python
t = Zeitstempel()
t.aktualisieren()
assert t.Datum == Zeitstempel().Datum, 'Initialize and update date!'
t = Kontostand(0.0)
assert t.Datum == Zeitstempel().Datum
t.einzahlen(10.0)
assert t.Datum == Zeitstempel().Datum
```

```python
a = Kontostand(0.0)
print(a)
```

```python
a.einzahlen(10.0)
print(a)
```

## Private Variablen
Wir erstellen eine Klasse Konto, in der der Name des Kontobesitzers sowie der Kontostand geschützte Attribute sind.

```python
class Konto():
    
    def __init__(self, name):
        self._Name = name
        self._Kontostand = Kontostand(0)
        
    def __str__(self):
        return ("Konto von %s:\n%s" % (self._Name, self._Kontostand.__str__()))
    
    def getKontostand(self):
        return self._Kontostand
    
    def einzahlen(self,x):
        self._Kontostand.einzahlen(x)
    def auszahlen(self,x):
        self._Kontostand.auszahlen(x)
        

k = Konto("Heiner Giefers")
k._Kontostand.einzahlen(0.22) # schlechter Stil
k.einzahlen(22.0) # schon besser
print(k) 
```

```python
class Konto():
    
    def __init__(self, name):
        self._Name = name
        self.__Kontostand = Kontostand(0) # Jetzt mit 2 Unterstrichen '__' !!!
        
    def __str__(self):
        return ("Konto von %s:\n%s" % (self._Name, self.__Kontostand.__str__()))
    
    def getKontostand(self):
        return self._Kontostand
    
    def einzahlen(self,x):
        self.__Kontostand.einzahlen(x)
    def auszahlen(self,x):
        self.__Kontostand.auszahlen(x)
        

k = Konto("Heiner Giefers")
k.__Kontostand.einzahlen(0.22) # schlechter Stil
k.einzahlen(22.0) # schon besser
print(k) 
```